# Assignment 1

Files and CSV

 The textfile, `assets/travel_plans.txt`, contains the summer travel plans for someone with some commentary. Find the total number of characters in the file and save to the variable `num`.

<!-- grade_id: cell-774c27f815c07a1e, points: 0, grade: false, solution: true -->
```python
num = None

# YOUR CODE HERE
inputfile1=open("assets/travel_plans.txt","r")
reader1=inputfile1.readlines()
for i in reader1:
    if num is None:
        num=len(i)
    else:
        num+=len(i)

#raise NotImplementedError()
inputfile1.close()
```

Sometimes, as in the next cell, we will show you one or more test cases, in the form of assert statements. Make sure your code pass those tests before you click the "Submit Assignment" button.

When you click the "Submit Assignment" button, there may be additional, hidden tests that are run. If you fail one of those tests, you will get feedback and can resubmit, as many times as you want, until you get them all right.

But it takes a little while for the auto grader to run. You will find it is faster, less frustrating, and more educational if you create additional tests of your own, either with assert statements or just print statements. That will give you a faster feedback-and-fix cycle; and formulating those tests will help you really understand what's going on. 

<!-- grade_id: cell-05ae8996efdc7e7b, points: 1, grade: true, solution: false -->
```python
assert num == 316, "num is not assigned the correct value"

```

We have provided a file called `assets/emotion_words.txt` that contains lines of words that describe emotions. Find the total number of words in the file and assign this value to the variable `num_words`.

<!-- grade_id: cell-0e1c86fffdf89e01, points: 0, grade: false, solution: true -->
```python
num_words = None

inputfile2=open("assets/emotion_words.txt","r")
reader2=inputfile2.readlines()
for i in reader2:
    if num_words is None:
        num_words=len(i.strip().split())
    else:
        num_words+=len(i.strip().split())
inputfile2.close()
#raise NotImplementedError()
```

<!-- grade_id: cell-7897d83bc4e61c17, points: 1, grade: true, solution: false -->
```python
assert num_words == 48, "num_words is not assigned the correct value"
```

Assign to the variable `num_lines` the number of lines in the file `assets/school_prompt.txt`.

**Note: different file than the previous question!**

<!-- grade_id: cell-75b61aedfeb412a9, points: 0, grade: false, solution: true -->
```python
num_lines = None

inputfile3=open("assets/school_prompt.txt","r")
reader3=inputfile3.readlines()
for i in reader3:
    if num_lines is None:
        num_lines=1
    else:
        num_lines+=1
inputfile3.close()
#raise NotImplementedError()
```

<!-- grade_id: cell-9122382cacdf4dda, points: 1, grade: true, solution: false -->
```python
assert num_lines == 10, "num_lines is not assigned the correct value"
```

Assign the first 30 characters of `assets/school_prompt.txt` as a string to the variable `beginning_chars`.

<!-- grade_id: cell-4836ee01c51ef411, points: 0, grade: false, solution: true -->
```python
beginning_chars = None

inputfile4=open("assets/school_prompt.txt","r")
reader4=inputfile4.read()
if beginning_chars is None:
    beginning_chars = "{}".format(reader4[:30])
inputfile4.close()
#raise NotImplementedError()
```

<!-- grade_id: cell-dae2a8b130911b15, points: 1, grade: true, solution: false -->
```python
assert type(beginning_chars) == type(""), "beginning_chars is not a string"
assert len(beginning_chars) == 30, "beginning_chars is not assigned the correct length"
```

**Challenge:** Using the file `assets/school_prompt.txt`, assign the third word of every line to a list called `three`.

<!-- grade_id: cell-3e6fa54e4fe79a87, points: 0, grade: false, solution: true -->
```python
three = None

inputfile = open("assets/school_prompt.txt", "r")
lines = inputfile.readlines()
for line in lines:
    words=line.split()
    if len(words) >= 3:  
        if three is None:
            three = []  
        three.append(words[2])
inputfile.close()
#raise NotImplementedError()
```

<!-- grade_id: cell-d7416a181aa2a896, points: 1, grade: true, solution: false -->
```python
assert type(three) == type([]), "three is not a list"
assert len(three) == 10, "three is not assigned the correct length"
```

**Challenge:** Create a list called `emotions` that contains the first word of every line in `assets/emotion_words.txt`.

<!-- grade_id: cell-cc5eb2cda7bc9765, points: 0, grade: false, solution: true -->
```python
emotions = None
with open("assets/emotion_words.txt","r") as file:
    lines=file.readlines()
    for line in lines:
        words=line.split()
        if len(words) > 0:
            if emotions is None:
                emotions = []  
            emotions.append(words[0])
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-cfad5425c19f0642, points: 1, grade: true, solution: false -->
```python
assert type(emotions) == type([]), "emotions is not a list"
assert len(emotions) == 7, "emotions is not assigned the correct length"
```

Assign the first 33 characters from the textfile, `assets/travel_plans.txt` to the variable `first_chars`.

<!-- grade_id: cell-fdaec3f02923071d, points: 0, grade: false, solution: true -->
```python
first_chars = None  
with open("assets/travel_plans.txt", "r") as file:
    content = file.read()  
    if first_chars is None: 
        first_chars = content[:33] 

#raise NotImplementedError()
```

<!-- grade_id: cell-29e19e35ea6223bb, points: 1, grade: true, solution: false -->
```python
assert type(first_chars) == type(""), "first_chars is not a string"
assert len(first_chars) == 33, "first_chars is not assigned the correct length"
```

**Challenge:** Using the file `assets/school_prompt.txt`, if the character ‘p’ is in a word, then add the word to a list called `p_words`.

<!-- grade_id: cell-20166062c59b6813, points: 0, grade: false, solution: true -->
```python
p_words = None 
with open("assets/school_prompt.txt", "r") as file:
    lines = file.readlines()  # Read all lines of the file
    for line in lines:  
        words = line.split()
        for word in words:  
            if 'p' in word:  
                if p_words is None:  
                    p_words = []
                p_words.append(word)  
#raise NotImplementedError()
```

<!-- grade_id: cell-fcfd04c36cdaf0a5, points: 1, grade: true, solution: false -->
```python
assert type(p_words) == type([]), "p_words is not a list"
assert len(p_words) == 5, "p_words is not assigned the correct length"
```

```python

```